// tests2.rs
// This test has a problem with it -- make the test compile! Make the test
// pass! Make the test fail! Execute `rustlings hint tests2` for hints :)

// I AM NOT DONE

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert_eq() {
        assert_eq!(1,1);

        // Each of the following are also valid solutions for this exercise

        //assert_eq!(true,true);
        //assert_eq!(false,false);
        //assert_eq!("b","b");

    }
}
